# Podstawy Sztucznej Inteligencji 2020/2021


Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

---

## Zaimportuj moduł  numpy  do przestrzeni nazw `np` 

<!-- grade_id: cell-39601c63625340b9, points: 0, grade: false, solution: true -->
```python
import numpy as np
```

<!-- grade_id: cell-657377bf7f83c7dc, points: 1, grade: true, solution: false -->
```python
from types import ModuleType
assert type(globals()['np']) == ModuleType
```

## Stwórz wektor zer o rozmiarze 12

<!-- grade_id: cell-2b8268ef9339c6a4, points: 0, grade: false, solution: true -->
```python
v = np.zeros(12)
```

```python
v
```

Output:
```
array([0., 0., 0., 0., 0., 0., 0., 0., 0., 0., 0., 0.])
```

<!-- grade_id: cell-b49cc2894c69dc72, points: 1, grade: true, solution: false -->
```python
assert v.sum() == 0
assert v.shape == (12,)
```

##  Napisz funkcję zwracającą ilość bajtów jaką zajmują dane w tablicy numpy

<!-- grade_id: cell-db49e414bf0a8307, points: 0, grade: false, solution: true -->
```python
def memof(a):
    return a.nbytes
```

```python
memof(v)
```

Output:
```
96
```

<!-- grade_id: cell-32ec95290a0502af, points: 1, grade: true, solution: false -->
```python
assert memof(np.zeros((12,10),dtype=np.int32))==480
```

```python
print(memof(np.zeros((12,10),dtype=np.int32)))
```

Output:
```
480
```

## Napisz funkcję tworzącą wektor zer o rozmiarze `n` z jedynką na piątym miejscu

<!-- grade_id: cell-d907a4380bbc262a, points: 0, grade: false, solution: true -->
```python
def make_v(n):
    v = np.zeros(n)
    v[4] = 1
    return v
```

```python
make_v(8)
```

Output:
```
array([0., 0., 0., 0., 1., 0., 0., 0.])
```

<!-- grade_id: cell-4dd989f0e98c6fcd, points: 1, grade: true, solution: false -->
```python
assert make_v(14)[4] == 1
np.testing.assert_array_equal( make_v(14)[:4] , 0)
np.testing.assert_array_equal( make_v(14)[5:] , 0)

```

## Utwórz wektor z wartościami od 10 do 29 włącznie

<!-- grade_id: cell-47d40a88c970b067, points: 0, grade: false, solution: true -->
```python
def make10_29():
    arr = np.arange(10, 30)
    return arr
```

```python
make10_29()
```

Output:
```
array([10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26,
       27, 28, 29])
```

<!-- grade_id: cell-8b23f3d048bd64fd, points: 1, grade: true, solution: false -->
```python
assert make10_29()[7] == 17
assert np.sum(make10_29()) == 390
```

## Odwóć kolejność elementów wektora

<!-- grade_id: cell-8daeafbeb3fabd80, points: 0, grade: false, solution: true -->
```python
def reverse(v):
    return np.flip(v)
```

```python
reverse(make_v(8))
```

Output:
```
array([0., 0., 0., 1., 0., 0., 0., 0.])
```

<!-- grade_id: cell-55510f70fe837a14, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal(reverse(np.array([1,2,3])), np.array([3,2,1]) )
```

## Stwórz macierz jednostkową o rozmiarze $n \times n$



<!-- grade_id: cell-92b16c324a05c06a, points: 0, grade: false, solution: true -->
```python
def identity_matrix(n=3):
    return np.identity(n)
```

```python
identity_matrix(5)
```

Output:
```
array([[1., 0., 0., 0., 0.],
       [0., 1., 0., 0., 0.],
       [0., 0., 1., 0., 0.],
       [0., 0., 0., 1., 0.],
       [0., 0., 0., 0., 1.]])
```

<!-- grade_id: cell-681b5089489167db, points: 1, grade: true, solution: false -->
```python
np.testing.assert_equal( identity_matrix(2), np.array([[1, 0],  [0, 1]]) )

```

## Utwórz macierz $n\times n$  z  wartościami od $1$ do $n^2$

 - tak by każdy rząd zawierał wartości rosnące o 1
 - tak by każda kolumna zawierała wartości rosnące o 1

<!-- grade_id: cell-8bb544fb4539f54e, points: 0, grade: false, solution: true -->
```python
def n2_col(n=3):
    col = np.arange(1, n**2 + 1).reshape(n, n)
    return col

def n2_row(n=3):
    row = np.arange(1, n**2 + 1).reshape(n, n)
    return row.T
```

```python
# test
print(n2_col())
print()
print(n2_row())
```

Output:
```
[[1 2 3]
 [4 5 6]
 [7 8 9]]

[[1 4 7]
 [2 5 8]
 [3 6 9]]
```

<!-- grade_id: cell-3eb2839edd26cc52, points: 2, grade: true, solution: false -->
```python
np.testing.assert_equal( n2_col(2), np.array([[1, 2],  [3, 4]]) )
np.testing.assert_equal( n2_row(2), np.array([[1, 3],  [2, 4]]) )

```

## Utwórz macierz $n\times m$  z  losowymi wartościami 

- z przedziału [0,1)
- z przedziału [a,b)

<!-- grade_id: cell-dc9ed83288917413, points: 0, grade: false, solution: true -->
```python
def rand1(n=3, m=4):
    return np.random.random(size=(n, m))

def rand2(n=3, m=4, a=6, b=7):
    return np.random.randint(a, b, size=(n, m))
```

```python
# test
print(rand1())
print()
print(rand2(a=1, b=10))
```

Output:
```
[[0.78070242 0.75128074 0.65489296 0.39705577]
 [0.84024785 0.12956789 0.89721097 0.47924681]
 [0.15590305 0.20334754 0.81861061 0.93587942]]

[[2 1 6 5]
 [7 4 3 5]
 [6 6 1 3]]
```

<!-- grade_id: cell-4748bfd79b06da2a, points: 2, grade: true, solution: false -->
```python
assert np.min(rand1(10, 11)) >= 0
assert np.max(rand1(10, 10)) < 1
assert np.min(rand2(10, 12, 3,7)) >= 3
assert np.max(rand2(10, 12, 3,7)) < 7
assert np.shape(rand2(10, 12, 3,7)) == (10, 12)
assert np.shape(rand1(10, 15)) == (10, 15)
```

## Znajdź indeksy dla których wartości wektora są równe zero

<!-- grade_id: cell-3dec5ac38532d35e, points: 0, grade: false, solution: true -->
```python
def is_zero(x):
    zeros = np.where(x == 0)[0]
    return zeros
```

<!-- grade_id: cell-52fa129f65194a82, points: 2, grade: true, solution: false -->
```python
x = np.array([1,2,0,1,0,11])
np.testing.assert_equal(is_zero(x), np.array([2,4]) ) 
```

```python
is_zero(x)
```

Output:
```
array([2, 4])
```

```python
def is_zero_v2(x):
    return np.argwhere(x == 0)
```

```python
is_zero_v2(x)
```

Output:
```
array([[2],
       [4]])
```

```python
def is_zero_true_or_false(x):
    return x != 0

print(is_zero_trueorfalse(x))
```

Output:
```
[ True  True False  True False  True]
```

## Oblicz dla zadanego wektora jego wartość najmniejszą, największą oraz średnią.

<!-- grade_id: cell-c15c1bf47a13ebad, points: 0, grade: false, solution: true -->
```python
def mystats(x):
    return np.amin(x), np.amax(x), np.mean(x)
```

<!-- grade_id: cell-6789a38d4549443a, points: 2, grade: true, solution: false -->
```python
x = np.array([1,2,0,1,0,11])
assert  mystats(x) == (0, 11, 2.5) 
```

```python
mystats(x)
```

Output:
```
(0, 11, 2.5)
```

## Stwórz dwuwymiarową tablicę z zerami w środku i jedynkami na zewnątrz

<!-- grade_id: cell-e2ca9f36bc360a8a, points: 0, grade: false, solution: true -->
```python
def zeros_padded(n=4):
    arr = np.zeros((n-2, n-2))
    arr_padded = np.pad(arr, 1, mode='constant', constant_values=1)
    return arr_padded
```

```python
zeros_padded(8)
```

Output:
```
array([[1., 1., 1., 1., 1., 1., 1., 1.],
       [1., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 1.],
       [1., 1., 1., 1., 1., 1., 1., 1.]])
```

<!-- grade_id: cell-cd548f3f78efd6cd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_equal(zeros_padded(4) , 
np.array([[ 1.,  1.,  1.,  1.],
         [ 1.,  0.,  0.,  1.],
         [ 1.,  0.,  0.,  1.],
         [ 1.,  1.,  1.,  1.]]))

```


## Używając `np.pad` dodaj do tablicy otoczenie z wartością 3

<!-- grade_id: cell-974a225dc7e83feb, points: 0, grade: false, solution: true -->
```python
def pad3(x):
    return np.pad(x, 1, mode='constant', constant_values=3)
```

```python
x = np.ones((6, 3))
pad3(x)
```

Output:
```
array([[3., 3., 3., 3., 3.],
       [3., 1., 1., 1., 3.],
       [3., 1., 1., 1., 3.],
       [3., 1., 1., 1., 3.],
       [3., 1., 1., 1., 3.],
       [3., 1., 1., 1., 3.],
       [3., 1., 1., 1., 3.],
       [3., 3., 3., 3., 3.]])
```

<!-- grade_id: cell-bce6102cad485520, points: 2, grade: true, solution: false -->
```python
x = np.ones((2,3))
x_pad3 = np.array([[ 3.,  3.,  3.,  3.,  3.],
                   [ 3.,  1.,  1.,  1.,  3.],
                   [ 3.,  1.,  1.,  1.,  3.],
                   [ 3.,  3.,  3.,  3.,  3.]])            
np.testing.assert_equal( pad3(x),x_pad3)
```

## Dla danej tablicy zastąp maksymalne wartości zerami


<!-- grade_id: cell-0af856baae027aa7, points: 0, grade: false, solution: true -->
```python
def maxto0(x):
    size = np.shape(x)
    #print(size)
    
    flat_x = x.flat[:]
    #print(flattened)
    
    flat_x = np.where(flat_x == np.max(flat_x), 0, flat_x)
    x = flat_x.reshape(size)
    
    return x
```

```python
x_expected = np.array([[1, 3, 1, 2],
                       [0, 2, 1, 0],
                       [2, 0, 2, 1]])

maxto0(x_expected)
```

Output:
```
array([[1, 0, 1, 2],
       [0, 2, 1, 0],
       [2, 0, 2, 1]])
```

```python
def maxto0_v2(x):
    temp = np.amax(x)
    #print(temp)
    return np.where(x == temp, 0, x)
```

```python
x_expected = np.array([[1, 3, 1, 2],
                       [0, 2, 1, 0],
                       [2, 0, 2, 1]])

maxto0_v2(x_expected)
```

Output:
```
array([[1, 0, 1, 2],
       [0, 2, 1, 0],
       [2, 0, 2, 1]])
```

<!-- grade_id: cell-9bd94b25f5056812, points: 2, grade: true, solution: false -->
```python
x_expected = np.array([[1, 0, 1, 2],
                       [0, 2, 1, 0],
                       [2, 0, 2, 1]])

x = np.array([[1, 3, 1, 2],
                   [3, 2, 1, 3],
                   [2, 0, 2, 1]])
np.testing.assert_equal(maxto0(x),x_expected)
```

## Broadcasting

Niech będzie dana tablica  $k$ parametrów  zmierzonych w  $n$  pomiarach: 

 - $ x_{ij} $: $j$-ty parametr w $i$-tym pomiarze

Oblicz korzystając z ``numpy`` nie używając pętli:

 - średnią po pomiarach dla wszyskich zmiennych: $$\mu_j = \frac{1}{n} \sum_{i=1}^n x_{ij}$$
 - odchylenie od wartości średniej dla każdej zmiennej we wszystkich pomiarach:$$\Delta_{ij} = x_{ij}-\mu_j$$
 - odchylenie średniokwadratowe dla każdej zmiennej: $$\sigma_j^2 = \frac{1}{n} \sum_{j=1}^n (x_{ij}-\mu_j)^2$$
 

<!-- grade_id: cell-3ba3e7c55658e301, points: 0, grade: false, solution: true -->
```python
def data_stats(x):
    x_avg = np.sum(x, axis=0) / x.shape[0]
    Delta = x - x_avg
    sigma = np.sum((x - x_avg)**2, axis=0) / x.shape[0]
    
    return x_avg, Delta, sigma
```

```python
x = np.array([[1, 3, 1, 2], [3, 2, 1, 3], [2, 0, 2, 1]])

data_stats(x)
```

Output:
```
(array([2.        , 1.66666667, 1.33333333, 2.        ]),
 array([[-1.        ,  1.33333333, -0.33333333,  0.        ],
        [ 1.        ,  0.33333333, -0.33333333,  1.        ],
        [ 0.        , -1.66666667,  0.66666667, -1.        ]]),
 array([0.66666667, 1.55555556, 0.22222222, 0.66666667]))
```

```python
data_stats(x)[0]  # x_avg
```

Output:
```
array([3.2, 3. , 2.2, 2.4])
```

```python
data_stats(x)[1]  # Delta
```

Output:
```
array([[-0.2,  1. , -0.2,  0.6],
       [-0.2,  1. ,  0.8,  1.6],
       [ 0.8,  0. , -0.2, -0.4],
       [-0.2, -1. , -1.2, -0.4],
       [-0.2, -1. ,  0.8, -1.4]])
```

```python
data_stats(x)[2]  # sigma
```

Output:
```
array([0.16, 0.8 , 0.56, 1.04])
```

<!-- grade_id: cell-6c8208eb3399b3a4, points: 4, grade: true, solution: false -->
```python
x = np.array([[3, 4, 2, 3],
       [3, 4, 3, 4],
       [4, 3, 2, 2],
       [3, 2, 1, 2],
       [3, 2, 3, 1]])
x_avg,Delta,sigma = data_stats(x)


np.testing.assert_allclose(x_avg,np.array([ 3.2,  3. ,  2.2,  2.4]))
np.testing.assert_allclose(Delta,\
  np.array([[-0.2,  1. , -0.2,  0.6],
           [-0.2,  1. ,  0.8,  1.6],
           [ 0.8,  0. , -0.2, -0.4],
           [-0.2, -1. , -1.2, -0.4],
           [-0.2, -1. ,  0.8, -1.4]]) )                           
np.testing.assert_allclose(sigma,np.array([ 0.16,  0.8 ,  0.56,  1.04]))

import inspect
blacklist = [".mean",".average","for","while","std",'var']
assert all([ not keyword  in inspect.getsource(data_stats) for keyword in blacklist])

```