# Podstawy Sztucznej Inteligencji 2020/2021


Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

---

# Naiwny Bayes

Notatnik ten zawiera implementacją naiwnego klasyfikatowa Bayesa wykorzystując bibliotekę `numpy`. Działanie klasyfikatora jest porównane implementacja ``GaussianNB`` w bibliotece `sklearn`.

Z twierdzenia Bayesa mamy:

$$P(J=j|X=x)=\frac{P(X=x|J=j)P(J=j)}{\sum_{j\in\{1,\ldots L\}}P(X=x|J=j)P(J=j)}$$

$$ p_j(x)=\frac{p_jf_j(x)}{f(x)} $$

Wykorzystamy te formuły w budowie klasyfikatora [Bayesa](http://books.icse.us.edu.pl/runestone/static/ai/KlasyfikacjaWOparciuPodobienstwoDoWzorca/OptymalnyKlasyfikatorStatystyczny.html).

```python
import numpy as np 
from sklearn.naive_bayes import GaussianNB

from sklearn.model_selection import train_test_split
from sklearn import datasets 
# dataset import 
iris = datasets.load_iris()

gnb = GaussianNB()

X = iris.data

y = iris.target
y = np.array(y)

train, test, train_targets, test_targets = \
   train_test_split(X, y, test_size=0.31)

# train 
clf = gnb.fit(train, train_targets)

# Test
Z = clf.predict(test)
```

```python
train.shape,test.shape
```

Output:
```
((103, 4), (47, 4))
```

### Wyznacz unikalne klasy na zbiorze trenującym.

<!-- grade_id: cell-5dc805642a59d6d9, points: 0, grade: false, solution: true -->
```python
Nlabels = np.unique(train_targets)
```

<!-- grade_id: cell-74b49f95665f4f84, points: 1, grade: true, solution: false -->
```python
np.testing.assert_array_equal(Nlabels,np.array([0, 1, 2]))
```

### Jaki jest procent poprawych odpowiedzi?

Implementacja ``GaussianNB`` daje tę odpowiedź w następujący sposób. Spróbuj otrzymać tą liczbę korzystając wyłącznie z funkcjonalności ``numpy`` i/lub języka Python.

```python
clf.score(test,test_targets)
```

Output:
```
0.9787234042553191
```

<!-- grade_id: cell-9e002553d865b165, points: 0, grade: false, solution: true -->
```python
a = np.sum(test_targets == Z)
b = np.size(Z)
correct = a/b
```

<!-- grade_id: cell-c05e0f08e67034cc, points: 1, grade: true, solution: false -->
```python
assert correct==clf.score(test,test_targets)
```

### Które odpowiedzi są złe:

```python
bad_idx = np.where(Z!=test_targets)
bad_idx
```

Output:
```
(array([27]),)
```

### Parametry klasyfikatora

```python
clf.theta_
```

Output:
```
array([[5.01      , 3.42      , 1.48333333, 0.24      ],
       [5.9025    , 2.77      , 4.2425    , 1.3275    ],
       [6.41212121, 2.93939394, 5.4030303 , 2.0030303 ]])
```

```python
clf.sigma_
```

Output:
```
array([[0.12756667, 0.13493334, 0.02405556, 0.01306667],
       [0.23424375, 0.0871    , 0.22694375, 0.03749375],
       [0.34045914, 0.08541782, 0.24150597, 0.08271809]])
```

```python
ith = 2
np.mean(train[train_targets == ith],axis=0)
```

Output:
```
array([6.41212121, 2.93939394, 5.4030303 , 2.0030303 ])
```

```python
np.var(train[train_targets == ith],axis=0)
```

Output:
```
array([0.34045914, 0.08541781, 0.24150597, 0.08271809])
```

## Implementacja naiwnego klasyfikatora Bayesa

**1\. Oblicz częstość występowania poszczególnych klasy $j$ w zbiorze treningowym $p_{j}$**

Jest to prawdopodobieństwo *a priori*.

<!-- grade_id: cell-05c1a043d5112990, points: 0, grade: false, solution: true -->
```python
p = np.unique(train_targets, return_counts = True)[1]/train_targets.shape[0]
#druga metoda:
#p = [np.sum(train_targets == ith)/train_targets.shape[0] for ith in range(3)]
```

<!-- grade_id: cell-c380329820017a6d, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(p,clf.class_prior_)
```

**2\. Oblicz wartość średnią dla każdej cechy z każdej klasy.**

Niech $\mu_{ij}$ oznacza  wartość średnią dla $j$-tej zmiennej w $i$-tej klasie, wtedy:
   
$$ \mu_{ij} =  \langle x_j \rangle_{ \forall x_j \in \;\mathrm{label} \;
i } $$

<!-- grade_id: cell-19b53a82e4d25284, points: 0, grade: false, solution: true -->
```python
mu = np.stack([np.mean(train[train_targets == ith], axis=0) for ith in range(3)])
```

<!-- grade_id: cell-cb6b4c5d617d34cd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(mu,clf.theta_)
```

```python
mu.shape
```

Output:
```
(3, 4)
```

**3\. Oblicz wariancję $j$-tej zmiennej w $i$-tej klasie.**

   $$ \sigma_{ij}^2 = \mathrm{Var} [x_j] _ {\;\;{ \forall x_j \in \mathrm{class}\; i}} $$

<!-- grade_id: cell-44ba370db5358650, points: 0, grade: false, solution: true -->
```python
sigma2 = np.stack([np.var(train[train_targets == ith], axis=0) for ith in range(3)])
```

<!-- grade_id: cell-2d5307ed0a70d2fd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(sigma2,clf.sigma_)
```

**4\. Oblicz prawdopodobieństwo *a posteriori* klasy $i$ dla danego wektora zmiennych  $\mathbf{x}$.**

Niech $k$ oznacza liczbę cech (zmiennych). W naszym przypadku mamy $k=4$.

Dla $i$-tej klasy mamy:


$$ P_i(\mathbf{x}) \simeq p_i f_i = p_i \frac{1}{\sqrt{(2\pi)^k\Pi_{j=1}^k\sigma_{ij}^2} } e^{ -\displaystyle\sum_{j=1}^{k}\frac{(x_j-\mu_{ij})^2}{2\sigma_{ij}^2} }
 $$
 
By otrzymać prawdopodobieństwa należy unormować $P_i$ dla każdego przypadku tak by suma $\sum_i P_i(\mathbf{x})=1$

<!-- grade_id: cell-d5a9a351f9292775, points: 0, grade: false, solution: true -->
```python
x  = test
P = None
k = test.shape[1] 

print ('number of features: k=',k)
#OK!!!
P = np.stack([p[i]*(1/(np.sqrt((2*np.pi)**k*np.prod(sigma2[i]))))*np.exp(-1/2.0*np.sum((x-mu[i])*(x-mu[i])/sigma2[i], axis=1)) for i in range(3)]).T
P = P/np.sum(P, axis = 1)[:, np.newaxis]

#???
#P = p/np.sqrt((2*np.pi)**k * np.prod(sigma2[i], axis=1))*np.exp(-0.5*np.sum(x[:, np.newaxis, :] - mu[np.newaxis, :, :])**2/sigma2[np.newaxis, :, :], axis=2)
#P = P/np.sum(P, axis=1)[:,newaxis]

```

Output:
```
number of features: k= 4
```

```python
sigma2.shape,x.shape,P.shape,np.unique(test_targets)
```

Output:
```
((3, 4), (47, 4), (47, 3), array([0, 1, 2]))
```

<!-- grade_id: cell-8890f9bbc94e3376, points: 4, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(P,clf.predict_proba(x))
```

**5\. Wyznacz klasę dla której prawdopodobieństwo *a posteriori* jest największe** 

 $i$ : $\quad\textrm{ gdy } P_i(x)=\max_{1\leq j\leq L} P_j(x)$


<!-- grade_id: cell-c19c296ed9d5d9be, points: 0, grade: false, solution: true -->
```python
prediction = np.argmax(P, axis=1)
```

<!-- grade_id: cell-8cfa7cdc5a2ae1d9, points: 3, grade: true, solution: false -->
```python
np.testing.assert_equal( prediction, clf.predict(test))
```

```python
print(clf.predict(test) )
print(prediction)
```

Output:
```
[1 0 2 2 0 0 0 0 2 0 0 2 0 2 2 0 2 2 0 2 0 0 2 2 1 0 1 2 2 0 1 1 1 2 1 0 2
 0 0 0 1 0 1 0 2 2 2]
[1 0 2 2 0 0 0 0 2 0 0 2 0 2 2 0 2 2 0 2 0 0 2 2 1 0 1 2 2 0 1 1 1 2 1 0 2
 0 0 0 1 0 1 0 2 2 2]
```

```python
clf.score(test,test_targets)
```

Output:
```
0.9787234042553191
```

```python
clf.predict(test) == prediction
```

Output:
```
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True])
```

```python

```