# Podstawy Sztucznej Inteligencji 2020/2021


Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

---

# Warstwowa sieć neuronowa

Rozważmy 3 warstwową sieć typu "fully connected".
Niech liczba neuronów w warstwach wynosi 

 - 3 na wejściu
 - 4 na pierwszej warstwie
 - 2 na drugiej warstwie
 - 3 na ostatnie warstwie (czyli wyjście jest 3 liczbami)
 
 

$$
\sigma(x) = \frac{1}{1+e^{-x}}
$$

 - pierwsza warstwa: $y_1 = \sigma( X\cdot W^1 + b^1)$
 - druga warstwa: $y_2 = \sigma( y_1\cdot W^2 + b^2)$
 - trzecia warstwa: $y_3 = \sigma( y_2\cdot W^3 + b^3)$
 
 
Zaimplementuj powyższe obliczenia używając zapisu macierzowego w numpy:

`np.dot(x,W)+b` - odpowiada macierzowemu mnożeniu:  $x\cdot W+b$, gdzie

 - $x$ jest wejściem - macierzą  wymiarze (shape) $n_{examples}\times n_{inputs} $ 
 - $W$ jest macierzą wag, każda kolumny odpowiadają neuronom: $n_{inputs} \times n_{neurons}$
 - $b$ to bias, wektor o rozmiarze $n_{neurons}$ (po jednym dla każdego neuronu).
 
 


```python
import numpy as np 
```

### Zaimplementuj funkcję  aktywacji
$$
\sigma(x) = \frac{1}{1+e^{-x}}
$$

<!-- grade_id: cell-94e1bd122cf3dccc, points: 0, grade: false, solution: true -->
```python
def s(x):
    return 1 / (1 + (np.e)**-x) 

#np.exp(-x)
```

<!-- grade_id: cell-0e836b6facb680cc, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(s(1.23),
                           0.773818574,rtol=1e-4)
```

<!-- grade_id: cell-b616ad0931f8ca71, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(s(np.array([1,2,3])),
                           [0.73105858, 0.88079708, 0.95257413],rtol=1e-4)
```

### Zainicjalizuj wagi

Przypisz wagom z i-tej warstwy takie same wartości równe $0.1\cdot i$   a bias-owi $0.2\cdot i$.

<!-- grade_id: cell-846e36eb9d573cb2, points: 0, grade: false, solution: true -->
```python
W1, b1 = np.ones((3, 4)) * 0.1, np.ones(4) * 0.2 
W2, b2 = np.ones((4, 2)) * 0.2, np.ones(2) * 0.4
W3, b3 = np.ones((2, 3)) * 0.3, np.ones(3) * 0.6
```

<!-- grade_id: cell-6a4368616b5b59ee, points: 1, grade: true, solution: false -->
```python
assert W1.shape == (3,4)
assert W2.shape == (4,2)
assert W3.shape == (2,3)
assert b1.shape[-1] == 4
assert b2.shape[-1] == 2
assert b3.shape[-1] == 3
```

Na wejsciu mamy zestaw dwóch przykładów.
Implementacja powinna przetworzyć poprawnie następujące dane:
    

```python
X = np.array([ [1,2,3],[1,0,0]])
```

```python
print(X.shape)
```

Output:
```
(2, 3)
```

### Zaimplementuj sieć


<!-- grade_id: cell-7eaeccfa6742b05b, points: 0, grade: false, solution: true -->
```python
y1 = s(np.dot(X, W1) + b1)
y2 = s(np.dot(y1, W2) + b2)
y3 = s(np.dot(y2, W3) + b3)
```

<!-- grade_id: cell-39474c51afdda396, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(y3,[[0.73747926, 0.73747926, 0.73747926],
       [0.73527256, 0.73527256, 0.73527256]],rtol=1e-4)

y2.shape == (2,2)
y1.shape == (2,4)
```

Output:
```
True
```