# Podstawy Sztucznej Inteligencji 2020/2021


Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

---

# Gini index

Indeksem Giniego zbioru $X_{train}=\{(x_i,y_i), i=1,...,N, y_i \in\{1,...,K\}\}$
gdzie $x_i$ jest wektorem cech a $y_i$ klasą nazywamy wartość:

$$ Gini(X_{train})=1−\sum_{i=1}^K \left(\frac{N_i}{N}\right)^2$$
gdzie $N_i$ jest ilością elementów zbioru uczącego z klasy $i$.


Indeksem Giniego zbioru $X_{train}$ ze względu na test $t$ nazywamy średnią ważoną indeksów Giniego dla podzbiorów zbioru uczącego spełniających test $t$ w każdym z możliwych wyników tego testu $r$:

$$ Gini(X_{train},t)=\sum_r \frac{N(t,r)}{N}Gini(X^{t,r}_{train}),$$

gdzie $X^{t,r}_{train}$  jest podzbiorem zbioru uczącego złożonym z elementów które dla testu $t$ dają wartość $r$, natomiast $N(t,r)$ jest liczebnością tego podzbioru.


import numpy as np 

### Zadanie 1

Zaimplementuj funkcję  zwaracjącą indeks Giniego dla zbioru zawierającego wiele klas. Niech liczebości będą kolejno wartościami na liscie, np:

    Ni =  [1, 3, 10]
    
oznacza ze mamy:
 - jednego reprezentanta klasy pierwszej, 
 - trzech klasy drugiej 
 - dziesięciu klasy trzeciej.

```python
import numpy as np 
```

<!-- grade_id: cell-a914bbc498d4c05d, points: 0, grade: false, solution: true -->
```python
# Gini = lambda Ni:1-sum(ni**2/sum(Ni)**2 for ni in Ni)

def Gini(Ni): 
    Gini_ = 1 - np.sum([(k_/np.sum(Ni))**2 for k_ in Ni])
    return Gini_
```

<!-- grade_id: cell-62c7e949d31e2058, points: 2, grade: true, solution: false -->
```python
Ni = [0,90]
assert Gini(Ni) == 0 


Ni = [2,2]
np.testing.assert_allclose(Gini(Ni) , 0.5, rtol=1e-6)


Ni = 100*[2]
np.testing.assert_allclose(Gini(Ni) , 0.99, rtol=1e-6)
```

### Zadanie 2.

Oblicz współczynnik Giniego zbioru który zawiera 12 reprezentatnów klasy A i 22 klasy B. W tescie zbiór ten jest podzielony na zbiory `Ni1` i `Ni2`. Oblicz indeks Giniego względem tego testu w następujących przypadkach:



<!-- grade_id: cell-f9be313ff987a5d0, points: 0, grade: false, solution: true -->
```python
Ni = [12, 22]
```

<!-- grade_id: cell-e1de9545be4a3ae2, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gini(Ni) ,0.45674740484429066, rtol=1e-6)
```

<!-- grade_id: cell-348e9e507cb186d7, points: 0, grade: false, solution: true -->
```python
Ni1, Ni2 = [11,2], [1,20]
assert sum(Ni) == sum(Ni1)+sum(Ni2)

Gt = (Gini(Ni1) * sum(Ni1) + Gini(Ni2) * sum(Ni2)) / sum(Ni)

print(Gt)
```

Output:
```
0.15556992027580266
```

<!-- grade_id: cell-4a63a87bfdf8d3a2, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gt ,0.1555, rtol=1e-3)
```

<!-- grade_id: cell-444ec9ae74954ff8, points: 0, grade: false, solution: true -->
```python
Ni1,Ni2 = [6,11], [6,11]

assert sum(Ni) == sum(Ni1)+sum(Ni2)

Gt = (Gini(Ni1) * sum(Ni1) + Gini(Ni2) * sum(Ni2)) / sum(Ni)

print(Gt)


```

Output:
```
0.45674740484429055
```

<!-- grade_id: cell-3edb43dc9a9ea037, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gt ,0.45674740, rtol=1e-3)
```

<!-- grade_id: cell-ba370522c770968c, points: 0, grade: false, solution: true -->
```python
Ni1,Ni2 = [10,10], [7,7]
assert sum(Ni) == sum(Ni1)+sum(Ni2)
Gt = (Gini(Ni1) * sum(Ni1) + Gini(Ni2) * sum(Ni2)) / sum(Ni)

print(Gt)
```

Output:
```
0.5
```

<!-- grade_id: cell-4cf283f1c7a023a4, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose( Gt ,0.5, rtol=1e-3)
```